```python
from IPython.display import HTML

hide_me = ''
HTML('''<script>
code_show=true; 
function code_toggle() {
  if (code_show) {
    $('div.input').each(function(id) {
      el = $(this).find('.cm-variable:first');
      if (id == 0 || el.text() == 'hide_me') {
        $(this).hide();
      }
      el = $(this).find('.cm-comment:first')
      if (id == 0 || el.text() == '# hide_me') {
        $(this).hide();
      }
    });
    $('div.output_prompt').css('opacity', 0);
  } else {
    $('div.input').each(function(id) {
      $(this).show();
    });
    $('div.output_prompt').css('opacity', 1);
  }
  code_show = !code_show
} 
$( document ).ready(code_toggle);
</script>
<form action="javascript:code_toggle()"><input style="opacity:0" type="submit" value="Click here to toggle on/off the raw code."></form>''')
```

## Exercise

Solve the following equation using addition.

$$x + (-6) = 5$$

<!-- grade_id: cell-2e081926423a22fd, points: 0, grade: false, solution: true -->
```python
from sympy import symbols, Eq, simplify

def ex():
    x = symbols('x')

    eq1 = Eq(x + (-6), 5)
    
    # YOUR CODE HERE
    raise NotImplementedError()

    return(eq1)

ex()  # Test out results
```

<!-- grade_id: cell-8d40c6a389ad40d4, points: 0, grade: true, solution: false -->
```python
# hide_me
# Unit tests
import inspect

assert 'ex' in locals(), 'Please keep the function name as `ex`'
assert (simplify(ex()) == Eq(symbols('x'), 11)), 'The final answer for x is incorrect'
assert '.lhs' in inspect.getsource(ex), 'You should use the lhs method to modify the left of the equation'
assert '.lhs' in inspect.getsource(ex), 'You should use the rhs method to modify the right of the equation'
assert 'solve' not in inspect.getsource(ex), 'Do not use the `solve()` function to get the answer'
assert type(ex()) == Eq, 'Your function should return an SymPy equation'

print("Great job!")
```

```python

```